# Part 1: Forward Mode Automatic Differentiation

Forward mode AD can simply be implemented by defining a class to represent [dual numbers](https://en.wikipedia.org/wiki/Dual_number) which hold the value and its derivative. The following skeleton defines a dual number and implements multiplication.

__Tasks:__

- Addition (`__add__`) is incomplete - can you finish it?
- Can you also implement division (`__truediv__`), subtraction (`__sub__`) and power (`__pow__`)?

<!-- grade_id: cell-a645b47036f56c2e, points: 0, grade: false, solution: true -->
```python
import math

class DualNumber:
    def __init__(self, value, dvalue):
        self.value = value
        self.dvalue = dvalue

    def __str__(self):
        return str(self.value) + " + " + str(self.dvalue) + "ε"

    def __mul__(self, other):
        return DualNumber(self.value * other.value,
            self.dvalue * other.value + other.dvalue * self.value)

    def __add__(self, other):
        return DualNumber(self.value + other.value, self.dvalue + other.dvalue)

    def __sub__(self, other):
        return DualNumber(self.value - other.value, self.dvalue - other.dvalue)

    def __pow__(self,other):
        return DualNumber(self.value**other.value, other.value*self.value**(other.value-1)*self.dvalue + math.log(self.value)*self.value**other.value*other.dvalue)

    def __truediv__(self,other):
        return DualNumber(self.value/other.value, (self.dvalue/other.value) - (self.value*other.dvalue/(other.value**2)))
```

<!-- grade_id: cell-7923d5c39d574cd7, points: 3, grade: true, solution: false -->
```python
# Tests

DualNumber(1,0) + DualNumber(1,0) / DualNumber(1,0) - DualNumber(1,0)**DualNumber(1,0)

```

Output:
```
<__main__.DualNumber at 0x7d2a4269ec90>
```

## Implementing math functions

We also need to implement some core math functions. Here's the sine function for a dual number:

```python
def sin(x):
    return DualNumber(math.sin(x.value), math.cos(x.value)*x.dvalue)
```

__Task:__ can you implement the _cosine_ (`cos`), _tangent_ (`tan`), and _exponential_ (`exp`) functions in the code block below?

<!-- grade_id: cell-61b9462cb781a8e1, points: 0, grade: false, solution: true -->
```python
def cos(x):
    return DualNumber(math.cos(x.value), -math.sin(x.value)*x.dvalue)

def tan(x):
    return DualNumber(math.tan(x.value), x.dvalue/math.cos(x.value)**2)

def exp(x):
    return DualNumber(math.exp(x.value),x.dvalue*math.exp(x.value))
```

<!-- grade_id: cell-9055ec23e6e9ffb8, points: 3, grade: true, solution: false -->
```python
# Tests
assert cos(DualNumber(0,0)).value == 1
assert tan(DualNumber(0,0)).value == 0
assert exp(DualNumber(0,0)).value == 1

```

## Time to try it out

We're now in a position to try our implementation.

__Task:__

- Try running the following code to compute the value of the function $z=x\cdot y+sin(x)$ given $x=0.5$ and $y=4.2$, together with the derivative $\partial z/\partial x$ at that point.

<!-- grade_id: cell-12f7c8fa42784946, points: 0, grade: false, solution: true -->
```python
x = DualNumber(0.5, 1.0)
y = DualNumber(4.2, 0)
z = x * y + sin(x)

print("z value at x=0.5:", z.value)
print("dz/dx at x=0.5:  ", z.dvalue)
```

Output:
```
z value at x=0.5: 2.579425538604203
dz/dx at x=0.5:   5.077582561890373
```

__Task__: Differentiate the above function with respect to $x$ and write the symbolic derivatives in the following box. Verify the result computed above is correct by plugging-in the values into your symbolic gradient expression.

<!-- grade_id: cell-1ae4d0a0801a142a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

__Task:__ Now use the code block below to compute the derivative $\partial z/\partial y$ of the above expression (at the same point $x=0.5, y=4.2$ as above) and store the derivative in the variable `dzdy` (just the derivative, not the Dual Number). Verify by hand that the result is correct.

<!-- grade_id: cell-eab763bf1b218dc8, points: 0, grade: false, solution: true -->
```python
x = DualNumber(0.5, 0)
y = DualNumber(4.2, 1)
z = x * y + sin(x)

dzdy = z.dvalue
print("dz/dy at x=0.5, y=4.2:", dzdy)
```

Output:
```
dz/dy at x=0.5, y=4.2: 0.5
```

<!-- grade_id: cell-a2f9f0ca4b1ed763, points: 2, grade: true, solution: false -->
```python
#Tests
assert dzdy
assert type(dzdy) == float

```

__Task:__ Finally, use the code block below to experiment and test the other math functions and methods you created.

<!-- grade_id: cell-b71d4bb64928ee38, points: 0, grade: true, solution: true -->
```python
x = DualNumber(0.3,1)
y = DualNumber(22,0)

z = y**x + sin(x)*tan(x)-y*exp(x)+ x*y

print("z value at x=0.5:", z.value)
print("dz/dx at x=0.5:  ", z.dvalue)
```

Output:
```
z value at x=0.5: -20.477771228748487
dz/dx at x=0.5:   0.7356756166057998
```